# These tests are meant to go hand-in-hand with the rendered alert rule.
# Use github.com/scality/action-prom-render-test@python-renderer python module
#
# Render the alerts file with
# gen-alert render alerts.yaml

evaluation_interval: 1m

rule_files:
  - alerts.rendered.yaml

tests:
  - interval: 30s
    input_series:
      - series: up{namespace="zenko", pod="mongodb"}
        values: 1 1 1 0 0

    alert_rule_test:
      - alertname: MongoDbDown
        eval_time: 1m
        exp_alerts: []
      - alertname: MongoDbDown
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              severity: critical
            exp_annotations:
              summary: MongoDb down
              description: MongoDb down

  - interval: 1m
    input_series:
      - series: mongodb_mongod_replset_member_state{namespace="zenko",pod="mongodb-0",state="PRIMARY"}
        values: 1 _

    alert_rule_test:
      - alertname: NoPrimary
        eval_time: 1m
        exp_alerts: []
      - alertname: NoPrimary
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              namespace: zenko
              state: PRIMARY
              severity: critical
            exp_annotations:
              description: "MongoDb down"
              summary: MongoDb down



  - interval: 30s
    input_series:
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-0", name="mongodb-0", state="PRIMARY"}
        values: 1 1 1 1
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-0", name="mongodb-1", state="SECONDARY"}
        values: 1 1 1 1
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-0", name="mongodb-2", state="SECONDARY"}
        values: 1 1 0 0
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-1", name="mongodb-0", state="PRIMARY"}
        values: 1 1 1 1
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-1", name="mongodb-1", state="SECONDARY"}
        values: 1 1 1 1
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-1", name="mongodb-2", state="SECONDARY"}
        values: 1 1 0 0
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-2", name="mongodb-0", state="PRIMARY"}
        values: 1 1 _ _
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-2", name="mongodb-1", state="SECONDARY"}
        values: 1 1 _ _
      - series: mongodb_mongod_replset_member_health{namespace="zenko",pod="mongodb-2", name="mongodb-2", state="SECONDARY"}
        values: 1 1 _ _

    alert_rule_test:
      - alertname: UnhealthyMemberWarning
        eval_time: 1m
        exp_alerts: []
      - alertname: UnhealthyMemberWarning
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              name: mongodb-2
              severity: warning
              state: SECONDARY
            exp_annotations:
              description: Member mongodb-2 (SECONDARY) is not healthy
              summary: Unhealthy MongoDb member

  - interval: 30s
    input_series:
      - series: mongodb_mongod_replset_member_election_date{namespace="zenko", pod="mongodb", set="rs0"}
        values: 1+0x9 10+1x10

    alert_rule_test:
      - alertname: TooManyElectionsWarning
        eval_time: 9m
        exp_alerts: []
      - alertname: TooManyElectionsWarning
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: warning
              set: rs0
            exp_annotations:
              description: Number of elections is greater than 10 in 10m
              summary: Too many elections

  - interval: 1m
    input_series:
      - series: mongodb_mongod_replset_member_replication_lag{namespace="zenko",pod="mongodb-0", name="mongo-0"}
        values: 8 12 15
      - series: mongodb_mongod_replset_member_replication_lag{namespace="zenko",pod="mongodb-1", name="mongo-1"}
        values: 4 9 12

    alert_rule_test:
      - alertname: ReplicationLagWarning
        eval_time: 1m
        exp_alerts: []
      - alertname: ReplicationLagWarning
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              severity: warning
              name: mongo-0
            exp_annotations:
              description: Mongodb replication lag is more than 10s on mongo-0
              summary: MongoDB replication lag
      - alertname: ReplicationLagWarning
        eval_time: 3m
        exp_alerts:
          - exp_labels:
              severity: warning
              name: mongo-0
            exp_annotations:
              description: Mongodb replication lag is more than 10s on mongo-0
              summary: MongoDB replication lag
          - exp_labels:
              severity: warning
              name: mongo-1
            exp_annotations:
              description: Mongodb replication lag is more than 10s on mongo-1
              summary: MongoDB replication lag

  - interval: 1m
    input_series:
      - series: mongodb_connections{namespace="zenko",pod="mongodb-0",state="current"}
        values: 90 98 110 110

    alert_rule_test:
      - alertname: TooManyClientConnectionsWarning
        eval_time: 2m
        exp_alerts: []
      - alertname: TooManyClientConnectionsWarning
        eval_time: 4m
        exp_alerts:
          - exp_labels:
              severity: warning
            exp_annotations:
              description: Too many client connections
              summary: Too many MongoDB client connections

  - interval: 1m
    input_series:
      - series: kubelet_volume_stats_available_bytes{namespace="zenko",persistentvolumeclaim="datadir-mongodb-0"}
        values: 30-5x3
      - series: kubelet_volume_stats_capacity_bytes{namespace="zenko",persistentvolumeclaim="datadir-mongodb-0"}
        values: 100+0x3

    alert_rule_test:
      - alertname: RemainingDiskSpaceWarning
        eval_time: 2m
        exp_alerts: []
      - alertname: RemainingDiskSpaceWarning
        eval_time: 4m
        exp_alerts:
          - exp_labels:
              severity: warning
              namespace: zenko
              persistentvolumeclaim: datadir-mongodb-0
            exp_annotations:
              description: MongoDb has low disk space
              summary: MongoDb has low disk space


